```python

```

# MTH793P: Advanced Machine Learning, Semester B, 2023/2024

---

**Full Name:**

<!-- grade_id: cell-06aa6f3b13940e27, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

**Student ID:**


<!-- grade_id: cell-2013f8791d3f4700, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

---

### General instructions

First, enter your **Full Name** and **Student ID** in the fields above.

**Read the following information carefully before turning to the actual exercises.**

This final project assignment is scored on a scale of 100 marks and will contribute **60% of your final mark** for this module.

This project is due by **Friday, 31 May 2024, 17:00 GMT**.
Submissions received after this deadline will be treated in accordance with the College Regulations for late submissions: Up to seven days after the deadline, late submissions will incur a penalty of 5% of the total marks per 24 hours (or a fraction thereof). For example, a submission received on 1/6/2024 at 18:00 GMT (25 hours after the deadline) will result in a deduction of 10 marks. Any submission received more than 168 hours after the deadline will receive zero marks.

The lecturer will be available to answer reasonable questions about this assignment until **Friday, 24 May 2024, 17:00 GMT**. You can contact the lecturer by email ([o.bobrowski@qmul.ac.uk](mailto:o.bobrowski@qmul.ac.uk)).

You must use this Jupyter Notebook to answer all exercises. You are not allowed to remove any cells from this notebook. You can add Markdown or code cells as needed.

You must submit this Jupyter Notebook through QMPlus, using the submission system for this assignment. You cannot submit your attempt in any other way. Files sent by other means, including email, will not be considered.

You cannot submit any additional files apart from this Jupyter Notebook. The ZIP file you downloaded contains various auxiliary files that you will need for the exercises. However, there is no need to upload these auxiliary files. The original versions of these files will be merged into your submission automatically.


For **Markdown** exercises, **replace all occurrences of** `YOUR ANSWER HERE`  **by valid Markdown** to answer the question.

For **coding** exercises,
**replace all occurrences of**
```python
### YOUR CODE HERE
raise NotImplementedError()
```
**by valid Python code** to solve the problem.
In your solution, you can **only use** built-in Python functions and **functions from modules were imported for you in this template**.
Importing other modules is not allowed and will result in withdrawal of the corresponding marks.
If your code produces an error and you cannot fix it, write a comment to show that you are aware of the problem and indicate possible causes if you can.
You can get **partial credit**, so even if you cannot solve an exercise completely, try to answer as much as you can.


Finally, before you turn this project in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart Kernel) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All (Jupyter Notebook Version 6) or Run$\rightarrow$Run All Cells (Jupyter Notebook Version 7)). Please **do not change the file name** when you upload your Notebook to QMPlus.

<div class="alert alert-block alert-warning">
    Please answer all the questions below by writing your own Python code. Using or copy-pasting pre-existing code from any source is not permitted, and will be considered as <b>plagiarism</b> and addressed according to College Regulations. You are not allowed to collaborate with other students or ask for help from any other source, including, but not limited to, online forums or ChatGPT (or any other AI agent). By submitting your Notebook via QMPlus, you confirm that you have followed these rules and that the submitted code is your own work. Upon violation, you will fail this test with zero marks.</div>

---

### Coding style

#### Comments

Comments are an essential part of programming.
Use appropriate comments to explain the reasoning behind your code.

#### Variables & Functions

Use the variable and function names specified (**<font color=red>in red</font>**) in each exercise when writing your code. If you use other names, the testing of your code might fail, resulting in a reduced mark. For any additional variables/functions you define, choose meaningful, concise names.


It is perfectly fine if you define extra auxiliary functions to carry out a task. This is to say, even if the exercise only asks you to write a function `example_function`, you are free to add other functions like `example_function_aux`, `some_other_function`, ... and use them in your implementation of `example_function`.
Choose meaningful names for all auxiliary functions, too.

---

### Assessment criteria

The first thing we will do to evaluate your submission is to clear all outputs and run the entire notebook from the beginning to the end. It is your responsibility to ensure that the code produces **no errors**. We will not try to fix any errors when assessing your work as you have several weeks to develop this project. Instead, you will receive **zero marks for any exercises where the code cell raises an error**. If you run into an error and do not know how to fix it, comment out the critical line of code and any subsequent ones that depend on it. Add a remark indicating that you are aware of the problem and potential reasons.

If you do not know how to implement a certain operation in Python,
add a comment and explain in plain English (and in sufficient detail) what you would like to do.
You may receive partial credit for expedient ideas formulated in this way.

**10% of the project mark** will be awarded for **coding style**. This includes the use of appropriate variable names, suitable comments to explain complex sequences of operations, and general organisation.
This is to say, even if your code is fully functional and solves the exercise, you may lose marks if the reasoning is not explained properly, or the code is not well-organised.

<div class="alert alert-block alert-danger">

**Running time:** You will be asked to code and use several ML methods that we studied in the lectures. If not implemented correctly, or used with the wrong parameters, these can take a long time to run. When you submit your notebook, you should make sure that:
* No single box in this notebook takes **longer than 30 seconds** to run on your personal computer.
* The entire notebook runs from beginning to end in **less than 3 minutes** on your personal computer.

You should make sure that your submitted notebook follows these limitations. Failing to do so will amount to a penalty of **20 points**. <br> 
However, if prior to submission you realise this is a potential issue - you should **comment out** the boxes that slow your code down,  add an explanation as to what you think is wrong with your code, and suggest ways to improve it. You will get **partial credit** for doing so. In this case, and only in this case, we ask that you also run your entire notebook (before commenting the slow sections) on your computer prior to submission. Export the complete notebook as a **PDF**, and attach the PDF to your submission, together with the notebook.
</div>

---

### Admitted features and imports

You are free to use any Python concepts or features that are built into the language. Furthermore, you can use any of the packages imported below. 
**You must not import any additional modules in any of the code you write.**

```python
import numpy as np
from numpy.linalg import svd

import matplotlib.pyplot as plt

from skimage import data
from skimage import img_as_float
from sklearn.cluster import KMeans
from sklearn.preprocessing import StandardScaler

from numpy.testing import assert_equal
from numpy.testing import assert_array_equal
from numpy.testing import assert_almost_equal
from numpy.testing import assert_array_almost_equal
```

---

# Part I - Image Compression with K-Means (20 points)

In this part we will explore the compression of colour images using the K-Means algorithm.<br>
A 24-bit RGB image uses $256 = 2^8$ possible values per colour-channel (Red/Green/Blue),
so that, each pixel in the image requires $3\times 8 = 24$ bits of storage. <br>
We would like to approximate such an image with an $n$-bit RGB image, where $n < 24$.


How do we choose the $2^n$ colour-intensity values, so that the image looks similar to the original image? <br>Your task is to show that this can be done with k-means clustering. 

For this part, use the k-means implementation in [**sklearn.cluster.KMeans**](http://scikit-learn.org/stable/modules/generated/sklearn.cluster.KMeans.html) (see documentation and examples in the link).

---

The following box loads an image from the **skimage** package and stores it in the variable **<font color=red>image</font>**.<br>
Choose any of the images by setting/removing the comment.

```python
### IMAGE SELECTION ###
# image = data.astronaut()
image = data.rocket()
# image = data.colorwheel()
```

```python
# Converting the values to float, and presenting the image.
image = img_as_float(image)
plt.imshow(image)
plt.axis('off')
plt.tight_layout;
```

Use **KMeans** (imported from sklearn.cluster) to apply the K-means clustering algorithm to the image you loaded.<br>
Set the number of clusters to be $2^\text{nbits}$, where **nbits** is number of bits used to store each pixel. <br>
Repeat the process for $\text{nbits} = 1,2,6$. Note that each image is a matrix of 480x480 RGB values, which you have to turn into a vector of 230,400 values.
<br>
The results should be stored as:<br>
- **<font color='red'>centers_\<i\> </font>** - an array containing the centers of the clusters. 
- **<font color='red'>labels_\<i\></font>**  - an array containing the labels for each pixel.<br>
    
    Where **<font color='red'>\<i\></font>** should be replaced with **<font color='red'>1,2,6</font>**.
    
    
**NOTE:** When using KMeans - use the default input parameters (except for number of clusters). 

<!-- grade_id: cell-3d7ac725682e7b05, points: 0, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

<!-- grade_id: cell-e3901cf86709f653, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-11b0fa871ac481c5, points: 12, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

Create a new image, where each pixel value is replaced by its nearest K-means center.<br>
Make sure the values you store are **between 0 and 1** (you can use **np.clip** for that).

Finally, before we can show the results, we need to convert the vector back to an image matrix. <br>
Store the final result in a variable named **<font color='red'>image_comp_\<i\></font>**, where you should replace **<font color='red'>\<i\></font>** with  **<font color='red'>1,2,6</font>**.

<!-- grade_id: cell-d0680929929776f8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a35aa487c7fdb7f7, points: 5, grade: true, solution: false -->
```python
### DO NOT REMOVE THIS CELL ###
```

Visualise your results. Are they good?

```python
# Plotting the three results.
plt.figure(figsize=(15, 5))
plt.subplot(1, 3, 1)
plt.imshow(image_comp_1)
plt.axis('off')
plt.title('1 bit')

plt.subplot(1, 3, 2)
plt.imshow(image_comp_2)  
plt.axis('off')
plt.title('2 bit')

plt.subplot(1, 3, 3)
plt.imshow(image_comp_6)  
plt.axis('off')
plt.title('6 bit')
```


Repeat this experiment for all the images in the block marked as **### IMAGE SELECTION ###**.<br>
To select a new image, simply comment out the current selection, and remove the comment from the image you want to check. <br>

The output for **one** of the images is noticalby inferior compared to the others.
In the next box, write down which image is has the worst output, and explain in words why this happens. After you checked all the images, leave one of them selected, you can choose any of the images.

<!-- grade_id: cell-668f426116be68ef, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

---

# Part II - PCA for Eigenfaces (37 points)

In this part, we work with a dataset consisting of numerous faces and create a basis of so-called eigenfaces.<br>
The images are **192x168**, and are stored as vectors of length **32,256**.

This utility function will just let us convert an image from vector to matrix representation, so it can be showed on the screen.

```python
def vec2img(vec):
    return np.reshape(vec,(168,192)).T
```

We load the faces database.<br>
**<font color='red'>FACES</font>** - an array where each entry is a collection of images of a single person <br>
**<font color='red'>NPEOPLE</font>** - number of people in the list (should be 20)<br>
**<font color='red'>NFACES</font>** - number of images per person (should be 64)

```python
f = open('faces.npy','rb')
FACES = np.load(f)
NPEOPLE = len(FACES)
NFACES = 64
NR = int(np.sqrt(NFACES))
f.close()
```

---

### The Eigenfaces of a Single Person

First, we will examine the photos of a **single** person.<br>
**<font color='red'>PI</font>** - the index of the person we examine. <br>
**<font color='red'>X_person</font>** - the data matrix, containing all images of person PI, as columns.<br>

You can change **PI** as you wish, to experiment with the photos of different people.

We start by presenting all the photos. 

```python
PI = 1
```

```python
X_person =  FACES[PI]

plt.figure(figsize=(15,15))
for i in range(NFACES):
    im = vec2img(X_person[:,i])
    plt.subplot(NR,NR,i+1)
    plt.imshow(im,cmap='gray')
    plt.axis('off')
```

Next, we want to find the "**eigenfaces**", i.e., the directions of the principal components for this collection of images.<br>
When using PCA, we should do all the processing for the **centred** data, i.e., with mean 0.
Therefore, take the following steps:
1. Define **<font color='red'>M_person</font>** to be the mean of all images in **X_person** (should be a 32,256-dimensional vector).
1. Subtract **M_person** from all images. Place the result in **<font color='red'>XZ_person</font>**.
1. Find the left-singular vectors of **XZ_person** (i.e., the matrix U in the SVD). Place the result in **<font color='red'>U_person</font>**.

At the end of this process, the columns of **U_person** should represent the **eigenfaces**.<br>

**<font color=blue>IMORTANT:</font>** When using **svd** (from np.linalg), it is important to specify **full_matrices=False**. Otherwise, the algorithm aims to compute the full dimensional matrices, which are not needed and require a lot of memory.

<!-- grade_id: cell-b202b416ce9712cd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d1a1b498ed538d95, points: 15, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

We will present the mean face **M_person** along with the first 15 eigenfaces in **U_person**.

```python
plt.figure(figsize=(15,15))

plt.subplot(4,4,1)
plt.imshow(vec2img(M_person), cmap='gray')
plt.axis('off')
plt.title('mean')

for i in range(15):
    plt.subplot(4,4,i+2)
    plt.imshow(vec2img(U_person[:,i]), cmap='gray')
    plt.axis('off')
    plt.title('$u_{%i}$' % (i+1))
```

Next, we want to try to get a feeling for what some of the eigenfaces represent.<br>
We start by exploring the role of the first 2 eigenfaces.<br>
To do so, we compute the reconstruction of **XZ_person** based on the first two principal components. We then add back the mean **M_person**.<br> The result should be placed in **<font color='red'>X12</font>**.

<!-- grade_id: cell-5eeb5592b35b5964, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-215b609eff147d9a, points: 3, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

We present the results for all the images. <br>

```python
plt.figure(figsize=(15,15))
for i in range(NFACES):
    im = vec2img(X12[:,i])
    plt.subplot(NR,NR,i+1)
    plt.imshow(im,cmap='gray')
    plt.axis('off')
```

**<font color='blue'>Question:</font>** Can you notice the effect of these two eigenfaces? what do you think they represent?

<!-- grade_id: cell-7a4336275c4ae228, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Next, we will do the same but for the 3rd and 4th eigenfaces. Place the result in **<font color='red'>X34</font>**.

<!-- grade_id: cell-e74dc0aa3950dda6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4ba08823454cb18f, points: 3, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

```python
plt.figure(figsize=(15,15))
for i in range(NFACES):
    im = vec2img(X34[:,i])
    plt.subplot(NR,NR,i+1)
    plt.imshow(im,cmap='gray')
    plt.axis('off')
```

**<font color='blue'>Question:</font>** Can you notice the effect of these two eigenfaces? Is it the same or different than the previous two?

<!-- grade_id: cell-x7a4336275c4ae228, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

---

### The Eigenfaces of a Collection of People

Next, we will have a similar exercise, but instead of taking a single person, we will take **all** of them (together).

We place all the images in **<font color='red'>X_all</font>** and show a random collection of them.

```python
X_all = np.concatenate(FACES,1)
         
plt.figure(figsize=(15,15))
for i in range(NFACES):
    j = np.random.randint(NFACES*NPEOPLE)
    im = vec2img(X_all[:,j])
    plt.subplot(NR,NR,i+1)
    plt.imshow(im,cmap='gray')
    plt.axis('off')
```

Find the eigenfaces in this case. This step is similar to what we did for a single person.<br>
Here, in addition to storing the singular vectors in **<font color='red'>U_all</font>**, you should also save the **singular values**. <br>
Make sure the singular values are stored in **<font color='red'>S_all</font>**.

This calculation can be a bit slow, but shouldn't take more than 20 seconds.

<!-- grade_id: cell-42404fa4cb898070, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f4699283da621af0, points: 3, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

Present the mean and the first 15 eigenfaces.

```python
plt.figure(figsize=(15,15))

plt.subplot(4,4,1)
plt.imshow(vec2img(M_all), cmap='gray')
plt.axis('off')
plt.title('mean')

for i in range(15):
    plt.subplot(4,4,i+2)
    plt.imshow(vec2img(U_all[:,i]), cmap='gray')
    plt.axis('off')
    plt.title('$u_{%i}$' % (i+1))
```

In the next box we try to test the effect of the  eigenfaces in different way. <br>
For each eigenface **u<sub>i</sub>** we want to show the image **M + t$\cdot$u<sub>i</sub>**, where **t** is in $\big[-\frac{\sigma_i}{20},\frac{\sigma_i}{20}\big]$, and $\sigma_i$ is the corresponding singular value.<br>
The list **<font color='red'>I</font>** indicates which eigenfaces to examine. You can experiment with different indexes.

```python
I = [0,1,2,3,4,5]
NT = 11
SCALE = 20
```

Fill in the missing code.

<!-- grade_id: cell-9fd67d60131abb38, points: 0, grade: false, solution: true -->
```python
NI = len(I)
M = X_all.mean(1)

plt.figure(figsize=(15,10))
for i in range(NI):
    T = np.linspace(-S_all[I[i]]/SCALE,S_all[I[i]]/SCALE,NT)
    for j in range(NT):
        # YOUR CODE HERE
        raise NotImplementedError()
        plt.subplot(NI,NT, i*NT+j+1)
        plt.imshow(vec2img(im),cmap='gray')
        plt.axis('off')
```

<!-- grade_id: cell-46003f1e45573a2a, points: 2, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

**<font color='blue'>Question:</font>** Can you notice the effects of the different eigenfaces? 

<!-- grade_id: cell-9e467dec6e7848a7, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Take the **first** image in the dataset, and place it in **<font color=red>im_test_1</font>**. Using the eigenfaces in **U_all**, modify the image so that it is **<font color=green>light on the left</font>** and **<font color=green>dark on the right</font>**. Save the result in **<font color=red>im_left_1</font>**. Similalry, make a version that is light on the right and dark on the left, and store it in **<font color=red>im_right_1</font>**. 

Repeate the same steps for the **50th** image, and store the results in **<font color=red>im_test_50, im_left_50, im_right_50</font>**.

<!-- grade_id: cell-4bf4f57df5fcfe99, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Present the original and modified images.

```python
plt.figure(figsize=(15,10))
plt.subplot(2,3,1)
plt.imshow(vec2img(im_test_1),cmap='gray')
plt.axis('off')
plt.title('Original image')

plt.subplot(2,3,2)
plt.imshow(vec2img(im_left_1),cmap='gray')
plt.axis('off')
plt.title('Light on the left')

plt.subplot(2,3,3)
plt.imshow(vec2img(im_right_1),cmap='gray')
plt.axis('off')
plt.title('Light on the right')

plt.subplot(2,3,4)
plt.imshow(vec2img(im_test_50),cmap='gray')
plt.axis('off')
plt.title('Original image')

plt.subplot(2,3,5)
plt.imshow(vec2img(im_left_50),cmap='gray')
plt.axis('off')
plt.title('Light on the left')

plt.subplot(2,3,6)
plt.imshow(vec2img(im_right_50),cmap='gray')
plt.axis('off')
plt.title('Light on the right')
```

**<font color='blue'>Question:</font>** Does the results look good? Do you notice a difference between the images? Try to discuss this.

<!-- grade_id: cell-4c4d364ee7640afe, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

---

# Part III -  Matrix Completion (33 points)

In this part, we perform  matrix completion  using the **Singular Value Thresholding** (SVT) algorithm we saw in class.<br>
Given an incomplete matrix **M**, recall that our objective is to solve:
$$ X^* = \arg\min_{X} \tau\|X|_* + \frac{1}{2} \|X\|_F^2\quad\text{subject to}\quad P_{\Omega}(X)=P_\Omega(M),$$
where $\Omega$ represents the set of indexes where the values of **M** are **known**.

The SVT algorithm should have the following update iterations:
$$
\begin{split}
X^{(k+1)} &= D_\tau(Y^{(k)}),\\
Y^{(k+1)} &= Y^{(k)} + \beta P_\Omega(M-X^{(k+1)}).\end{split}
$$
The starting value should be $Y^{(0)} = 0$.<br> 
The algorithm should stop after $k$ steps if:
$$ \frac{\|P_\Omega(M)-P_\Omega(X^{(k)})\|_F}{\|P_\Omega(M)\|_F} \le \text{tolerance}.$$


Place your code in the function **<font color='red'>matrix_complete</font>**. It should take the following input:
* **<font color='red'>M</font>** - the incomplete matrix in $\mathbb{R}^{D\times N}$. The **known** values are assumed to be between 0 and 1, and an unknown value will be marked by **<font color='red'>-1</font>**.
* **<font color='red'>tolerance</font>** - the maximum tolerance to decide when iterations should stop. The default value is $8.1\times 10^{-5}$.
* **<font color='red'>beta_val</font>** - a value for the increment size $\beta$.<br>
If **beta_val** is not specified (=None), the default value should be $\beta = 0.82\times\frac{D\times N}{K}$, where **K** is the number of **known** entries in **M**.
* **<font color='red'>tau_val</font>** - a value for the shrinking size $\tau$. <br>
If **tau_val** is not specified (=None), the default value should be $\tau = 0.37\times\max(D,N)$.
* **<font color='red'>max_iter</font>** - maximum number of iterations to run.
* **<font color='red'>iter_print</font>** - how often to print out a progress line.

The output of the function should be the final value of **<font color='red'>X</font>**.



<!-- grade_id: cell-2f0b96cb789b0c25, points: 0, grade: false, solution: true -->
```python
### IF YOU NEED SOME AUXILIARY FUNCTIONS YOU CAN PLACE THEM HERE
### (ALSO OK TO LEAVE EMPTY IF YOU DON'T)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-41075ce2326f072e, points: 0, grade: false, solution: true -->
```python
def matrix_complete(M, tolerance=8.1e-5, beta_val=None, tau_val=None, max_iter=10000, iter_print=100):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a47e012856ac0e0d, points: 15, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###

np.random.seed(11232)
M_test = np.random.rand(4,5)
M_test[0,2] = -1
M_test[3,1] = -1
M_test[1,0] = -1
M_test[2,3] = -1

X_correct = np.array([[0.28184808, 0.22232085, 0.25328388, 0.50521773, 0.82542652],
       [0.21899948, 0.0900134 , 0.84216102, 0.65204837, 0.01802106],
       [0.07918562, 0.90582226, 0.65758943, 0.15359926, 0.48446993],
       [0.3112642 , 0.23217241, 0.88039408, 0.6651078 , 0.32746706]])

X_test = matrix_complete(M_test, tolerance=1e-4, beta_val = 0.5*20/16, tau_val=1.0)

assert_almost_equal(X_test, X_correct)


print('--- TEST PASSED ---')
```

The function **<font color='red'>matrix_corrupt</font>** takes a matrix **X**, assuming to have positive values, and with probability **P** replaces each element with a value of **<font color='red'>-1</font>** to mark it as a **missing value**.

```python
def matrix_corrupt(X, P):
    Xnew = np.array(X)
    R = np.random.rand(X.shape[0], X.shape[1])
    IDX = np.where(R<P)
    Xnew[IDX] = -1
    return Xnew
```

---

### RGB Images
We will use the matrix completion algorithm to correct missing values in RGB images.<br>
We will us the same photo selection as in Part I. Examine all the picture, and leave one selected when you submit your project (doesn't matter which).

```python
### IMAGE SELECTION ###
image = data.astronaut()
# image = data.rocket()
# image = data.colorwheel()
```

```python
# Converting the RGB values into the range [0,1]
image = img_as_float(image)
```

The missing values can appear in any of the R/G/B channels independently.<br> Therefore, if the image is of width  **W** and height **H**, we will process it as a matrix with **H rows** and **3W columns**.<br>
We start by writing functions to convert between the two formats.

You will need to code two functions:
* **<font color=red>im2mat</font>** - converts an image from an **(H x W x 3)** 3d-array to an **(H x 3W)** matrix.
* **<font color=red>mat2im</font>** - converts an **(H x 3W)** matrix back into an **(H x W x 3)** 3d-array.

<!-- grade_id: cell-c202341df77c1676, points: 0, grade: false, solution: true -->
```python
def im2mat(image):
    # YOUR CODE HERE
    raise NotImplementedError()

def mat2im(mat):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Testing the **im2mat** and **mat2im** functions.

<!-- grade_id: cell-a3c13eb40e67e60d, points: 2, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###

np.random.seed(2343234)

test_im = np.random.rand(2,3,3)
test_mat = im2mat(test_im)
test_im_2 = mat2im(test_mat)

test_mat_RES = np.array([[0.21435714, 0.9483089 , 0.26014289, 0.85474945, 0.3003837 ,
        0.86430052, 0.74878394, 0.44215536, 0.28218045],
       [0.75837569, 0.50741801, 0.90138123, 0.99696511, 0.55819116,
        0.91226562, 0.92116696, 0.96942254, 0.63157848]])

assert_array_almost_equal(test_mat, test_mat_RES)
assert_array_almost_equal(test_im_2, test_im)


print('--- TEST PASSED ---')
```

Next, convert the photo stored in **image** to a matrix using **im2mat**, save the result in **<font color=red>im_mat</font>**.

<!-- grade_id: cell-b69e08354211a7ac, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

If all was done well, presenting **im_mat** you should see 3 copies of the same photo.

```python
plt.figure(figsize=(15,5))
plt.imshow(im_mat, cmap='gray');
plt.axis('off')
```

Next, we will erase random values from the matrix, to simulate missing values.<br>
Use the function **matrix_corrupt** to generate a copy of **im_mat** with missing values. Use **P=0.3**, and save the result in **<font color=red>im_mat_cor</font>**. 

In addition, since we will want to show what the corrupt image looks like, we will have to convert the matrix version back to an RGB format. Store the RGB version in **<font color=red>im_cor</font>**. Note that when presenting the image, the values are expected to be between 0 and 1. Since our missing values are **-1** it will damage the view. Use **np.clip** to amend this (only for **im_cor**, do NOT do this for **im_mat_cor**).

<!-- grade_id: cell-6f65144d05827726, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-21ad749bcc4c4449, points: 3, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

The next step is to complete the missing value in the matrix **im_mat_cor**. To do that, use the function **matrix_complete** you coded earlier. Run it with $\beta=1.23$, $\tau=97$, and max_iter=50. Store the result in **<font color=red>im_mat_comp</font>**.

After the matrix completion is done, convert the matrix back into an image, and clip the values like you did earlier. Store the result in **<font color=red>im_comp</font>**. 

<!-- grade_id: cell-11492a72be2f7682, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-51b108bf207f9a5f, points: 5, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

```python
plt.figure(figsize=(15,5))
plt.subplot(1,3,1)
plt.imshow(image)
plt.axis('off')
plt.title('Original image')

plt.subplot(1,3,2)
plt.imshow(im_cor);
plt.axis('off');
plt.title('Corrupt image')

plt.subplot(1,3,3)
plt.imshow(im_comp)
plt.title('Completed image')
plt.axis('off');
```

The box below is free form.
Repeat the experiment above, but try to test the effect of the parameters **P**, **tolerance**, $\beta$, $\tau$, etc.
When you submit your final version, include 2 examples, and discuss the results in the box below. Make sure that the examples you choose demonstrate some differences that you can discuss. Use the same variable names, so they are properly displayed in the following box.

<!-- grade_id: cell-4c1d290dd17bffd3, points: 2, grade: true, solution: true -->
```python
### SETTING -1-
# YOUR CODE HERE
raise NotImplementedError()
```

```python
plt.figure(figsize=(15,5))
plt.subplot(1,3,1)
plt.imshow(image)
plt.axis('off')
plt.title('Original image')

plt.subplot(1,3,2)
plt.imshow(im_cor);
plt.axis('off');
plt.title('Corrupt image')

plt.subplot(1,3,3)
plt.imshow(im_comp)
plt.title('Completed image')
plt.axis('off');
```

<!-- grade_id: cell-f98598eaf6890ac2, points: 2, grade: true, solution: true -->
```python
### SETTING -2-
# YOUR CODE HERE
raise NotImplementedError()
```

```python
plt.figure(figsize=(15,5))
plt.subplot(1,3,1)
plt.imshow(image)
plt.axis('off')
plt.title('Original image')

plt.subplot(1,3,2)
plt.imshow(im_cor);
plt.axis('off');
plt.title('Corrupt image')

plt.subplot(1,3,3)
plt.imshow(im_comp)
plt.title('Completed image')
plt.axis('off');
```

Discuss the results here, in particular,  the effect of the different parameter values (e.g., on the quality of the output, speed of convergence, etc.), and possible reasons for these effects.

<!-- grade_id: cell-88bd25acbf50fa91, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

```python

```